## Exercise 01.1

Degrees Fahrenheit ($T_F$) are converted to degrees Celsius ($T_c$) using the formula:

$$
T_c = 5(T_f - 32)/9
$$

Write a program to convert 46 degrees Fahrenheit to degrees Celsius and print the result.
Write your program such that you can easily change the temperature in Fahrenheit that you are converting from.

Use the variable name `T_c` for the temperature in degrees Celsius.

<!-- grade_id: cell-ca2d6af7fabf2918, points: 0, grade: false, solution: true -->
```python
T_f = 46
...
```

<!-- grade_id: cell-a7542313c8cdf4ed, points: 0, grade: true, solution: false -->
```python
## test ##
import math
assert math.isclose(T_c, 7.77777777777777)
```

```python
# Temperature in Fahrenheit
F = int(46)

# Convert to Celsius
T_c = (5/9) * (F - 32)

# Print result
print("Temperature in Celsius:", T_c)

```

Output:
```
Temperature in Celsius: 7.777777777777779
```

## Exercise 01.2

You have been tasked with developing a mortgage affordability test tool. A component is to compute the monthly
interest payments for a range of scenarios.

Interest on a particular mortgage is charged at fixed rate above the Bank of England (BoE) 'official Bank Rate'.
Interest is computed per annum, and interest payments spread equally over each month of the year.
Write a program that computes the interest payable each month, with variables for:

1. Loan principal (amount borrowed, fixed)
1. Official Bank Rate (percentage, expressed per annum)
1. Rate over the official Bank Rate (percentage, expressed per annum)

Test your program with a loan principal of £172,000, BoE rate of 2.25%, and rate over the
BoE rate of 1.49%. Use the variable name `interest` for the monthly interest amount.

<!-- grade_id: cell-aaab53918bb9627d, points: 0, grade: false, solution: true -->
```python
# Assign data to variables
principal = 172000

...
```

<!-- grade_id: cell-258e3a2c80f6a765, points: 0, grade: true, solution: false -->
```python
## test ##
assert math.isclose(interest, 536.0666666666667)
```

```python
import math

# Assign data to variables
principal = 172000
boe_rate = 2.25
rate_over_boe = 1.49

# Total annual interest rate
annual_rate = (boe_rate + rate_over_boe) / 100

# Monthly interest
interest = principal * annual_rate / 12

print(interest)

## test ##
assert math.isclose(interest, 536.0666666666667)

```

Output:
```
536.0666666666667
```